# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    # YOUR CODE HERE
    k=[]
    l=[]
    if len(x)==len(y):
        for i in range(0, len(x)):
            k.append(int(x[i]))
        for j in range(0, len(y)):
            l.append(int(y[j]))
        for a in range(0, len(k)):
            if sum(k)== sum(l):
                return True
    else:
        return False
print(same_digits('501', '150'))
```

Output:
```
True
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```